# Petite mise en jambes programmation Python

Quelques questions dans lesquelles il faut définir des fonctions. Ces questions seront corrigées automatiquement par l'outil. Et une dernière où l'on demande d'écrire un programme qui utilise une des fonctions des questions précédentes et qui sera évaluée par le formateur.

## Question 1

Ecrire une fonction `partie_entiere` qui prend un entier `n` **positif ou nul** en paramètre et construit la liste des entiers 0 et 1 correspondant à l'écriture binaire de `n` par la méthode des divisions successives.

<!-- grade_id: partie_entiere_def, points: 0, grade: false, solution: true -->
```python
# Mettre votre code dans cette cellule

```

<!-- grade_id: partie_entiere_tests, points: 1, grade: true, solution: false -->
```python
"""Testez que votre fonction retourne les bons résultats pour quelques entrées.
    Cette cellule n'est pas modifiable, vous pouvez simplement l'exécuter."""

assert bin_partie_entiere(0) == [0]
assert bin_partie_entiere(12) == [1, 1, 0, 0]


```

## Question 2

Ecrire une fonction `bin_partie_decimale` qui prend un float f < 1 en paramètre et construit la liste des entiers 0 et 1 correspondant à l'écriture binaire de f par la méthode des multiplications successives. 

**Note :** Cette méthode, en théorie peut ne pas terminer puisqu'un décimal peut ne pas s'écrire comme une somme finie de 1/2^k. Dans notre fonction, comme nous utilisons un float le problème ne se posera pas et nous aurons une approximation seulement.

**Note 2 :** On renverra la liste vide dans le cas de 0

<!-- grade_id: partie_decimale_def, points: 0, grade: false, solution: true -->
```python
# Mettre votre code dans cette cellule
```

<!-- grade_id: partie_decimale_tests, points: 1, grade: true, solution: false -->
```python
"""Testez que votre fonction retourne les bons résultats pour quelques entrées.
    Cette cellule n'est pas modifiable, vous pouvez simplement l'exécuter."""

assert bin_partie_decimale(0) == []
assert bin_partie_decimale(.625) == [1, 0, 1]
assert bin_partie_decimale(.7) == [1, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1, 0, 0, 1, 1]

```

## Question 3

A l'aide des fonctions précédentes, écrire une fonction `binaire` qui calcule la chaîne de caractères correspondant à l'écriture binaire d'un nombre (entier ou à virgule). 

<!-- grade_id: binaire_def, points: 0, grade: false, solution: true -->
```python
# Mettre votre code dans cette cellule
```

<!-- grade_id: binaire_tests, points: 1, grade: true, solution: false -->
```python
"""Testez que votre fonction retourne les bons résultats pour quelques entrées.
    Cette cellule n'est pas modifiable, vous pouvez simplement l'exécuter."""

assert binaire(0) == '0'
assert binaire(251.0) == '11111011'
assert binaire(11.7) == '1011,101100110011001100110011001100110011001100110011'
```

## Question 4

Ecrire un programme qui demande un nombre à l'utilisateur et affiche l'écriture binaire de ce nombre.

<!-- grade_id: main_exo1, points: 1, grade: true, solution: true -->
```python
"""Complétez le code suivant"""
# Mettre votre code dans cette cellule

# instruction 1
try:
    # instruction 2
except:
    # instruction 3

# print(quelque_chose)
```