# Manipulation des itérateurs

Un petit exercice pour manipuler les itérateurs

## fonction `is_prime` 

Ecrire une fonction `is_prime` qui prend un entier n positif ou nul et retourne le booléen `True` si et seulement si n est un nombre premier.

<!-- grade_id: is_prime_v1_def, points: 0, grade: false, solution: true -->
```python
from math import sqrt

# Mettre votre code dans cette cellule
```

<!-- grade_id: is_prime_v1_tests, points: 1, grade: true, solution: false -->
```python
"""Validez votre fonction grâce aux exemples ci-dessous. Cette cellule est verrouillée
   Vous ne pouvez pas la modifier, juste l'exécuter"""

assert is_prime(0) == False
assert is_prime(2) == True
assert is_prime(19) == True
assert is_prime(100001357) == True
assert is_prime(22222222) == False


```