# Des exercices pour choisir les bonnes structures

Voici quelques exercices, pas difficiles au niveau algorithmique mais dont les solutions s'écrivent plus ou moins élégamment en fonction de la structure choisie, des méthodes prédéfinies utilisées avec ces structures. Liste ? Dictionnaire ? Chaînes de caractères ? A vous de choisir et de résoudre au mieux les petits exemples ci-dessous.

## Le problème de Joseph

Un grand classique. Des chevaliers sont en cercle. Ils vont se donner la mort. Le premier tue son voisin de gauche. Le voisin de gauche du chevalier tombé fait de même. Et ainsi de suite jusqu'à qu'il ne reste plus qu'un seul chevalier debout. (vous trouverez ici : [une vidéo youtube illustrant le problème](https://www.youtube.com/watch?v=uCsD3ZGzMgE)).

Vous devez écrire une fonction `chevaliers` qui prend un entier N (le nombre de chevaliers) en paramètre et qui calcule le numéro du chevalier survivant (on suppose qu'ils sont numérotés de 1 à N).

<!-- grade_id: joseph_def, points: 0, grade: false, solution: true -->
```python
# Mettre votre code dans cette cellule
```

<!-- grade_id: joseph_tests, points: 1, grade: true, solution: false -->
```python
"""Testez que votre fonction retourne les bons résultats pour quelques entrées.
    Cette cellule n'est pas modifiable, vous pouvez simplement l'exécuter."""

assert chevaliers(10) == 5
assert chevaliers(1) == 1
assert chevaliers(8) == 1
assert chevaliers(42) == 21
```

## Balade en centre ville

Une personne a rendez-vous au coin d'une rue (de deux pour exact) mais arrive avec quelques minutes d'avance. Elle décide de faire un petit tour, au hasard, laissant son smartphone générer une balade aléatoire. La ville est constituée de blocs carrés et les rues forment un quadrillage. L'application fournit donc une balade sous la forme d'une chaîne de caractères constituée de 'n' (pour la direction Nord), 's' (Sud), 'o' (Ouest) et 'e' (Est). Notre promeneur souhaite revenir à son point de rendez-vous **exactement** à la bonne heure. Sachant que longer un bloc prend 1 minute et que la balade est constituée d'un nombre entier de blocs, écrire une fonction `est_valide` qui valide un itinéraire suggéré par le smartphone si ce dernier satisfait la contrainte. Ci-dessous trois itinéraires, un vert et deux rouges pour une personne ayant 12 minutes d'avance. Le vert est valide alors que les deux autres ne le sont pas (un retour ailleurs qu'au point de rendez-vous et l'autre nécessite 14 minutes).

![parcours en ville](../../ressources_files/trois_parcours.png)

- La verte : `'nneeenossoos'`
- La rouge : `'ooonnnnneess'``
- La mauve : `'ssenenessooonn'`

<!-- grade_id: balade_def, points: 0, grade: false, solution: true -->
```python
# Mettre votre code dans cette cellule
```

<!-- grade_id: balade_tests, points: 1, grade: true, solution: false -->
```python
"""Testez que votre fonction retourne les bons résultats pour quelques entrées.
    Cette cellule n'est pas modifiable, vous pouvez simplement l'exécuter."""

assert est_valide('nneeenossoos', 12) == True
assert est_valide('ooonnnnneess', 12) == False
assert est_valide('ssenenessooonn', 12) == False
```

## Bien parenthésée

On considère un expression comportant du texte et des parenthèses, accolades, crochets. Il s'agit d'écrire une fonction `bien_parenthesee` qui prend une chaîne de caractères S et retourne `True` si et seulement si S est bien parenthésée. Voir la cellule de tests pour des exemples et contre-exemples.

<!-- grade_id: parentheses_def, points: 0, grade: false, solution: true -->
```python
# Mettre votre code dans cette cellule
```

<!-- grade_id: parentheses_tests, points: 1, grade: true, solution: false -->
```python
"""Testez que votre fonction retourne les bons résultats pour quelques entrées.
    Cette cellule n'est pas modifiable, vous pouvez simplement l'exécuter."""

assert bien_parenthesee('(([3-4]*12) + {5/3})') == True
assert bien_parenthesee('([})') == False, 'crochet ouvrant pour une accolade fermante'
assert bien_parenthesee('()()({()())') == False, 'manque une accolade fermante'
```